## Tips
- To avoid unpleasant surprises, I suggest you _run all cells in their order of appearance_ (__Cell__ $\rightarrow$ __Run All__).


- If the changes you've made to your solution don't seem to be showing up, try running __Kernel__ $\rightarrow$ __Restart & Run All__ from the menu.


- Before submitting your assignment, make sure everything runs as expected. First, restart the kernel (from the menu, select __Kernel__ $\rightarrow$ __Restart__) and then **run all cells** (from the menu, select __Cell__ $\rightarrow$ __Run All__).

## Reminder

- Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name, UA email, and collaborators below:



Several of the cells in this notebook are **read only** to ensure instructions aren't unintentionally altered.  

If you can't edit the cell, it is probably intentional.

```python
NAME = "Wade Swede"
# University of Arizona email address
EMAIL = "was2@email.arizona.edu"
# Names of any collaborators.  Write N/A if none.
COLLABORATORS = "N/A"
```

## Scratchpad

You are welcome to create new cells (see the __Cell__ menu) to experiment and debug your solution.

```python
%load_ext autoreload
%autoreload 2
```

Output:
```
The autoreload extension is already loaded. To reload it, use:
  %reload_ext autoreload
```

# Mini Python tutorial

This course uses Python 3.8.

Below is a very basic (and incomplete) overview of the Python language... 

For those completely new to Python, [this section of the official documentation may be useful](https://docs.python.org/3.8/library/stdtypes.html#common-sequence-operations).

```python
# This is a comment.  
# Any line starting with # will be interpreted as a comment

# this is a string assigned to a variable
greeting = "hello"

# If enclosed in triple quotes, strings can also be multiline:

"""
I'm a multiline
string.
"""

# let's use a for loop to print it letter by letter
for letter in greeting:
    print(letter)
    
# Did you notice the indentation there?  Whitespace matters in Python!

# here's a list of integers

numbers = [1, 2, 3, 4]

# let's add one to each number using a list comprehension
# and assign the result to a variable called res
# list comprehensions are used widely in Python (they're very Pythonic!)

res = [num + 1 for num in numbers]

# let's confirm that it worked
print(res)

# now let's try spicing things up using a conditional to filter out all values greater than or equal to 3...
print([num for num in res if not num >= 3])

# Python 3.7 introduced "f-strings" as a convenient way of formatting strings using templates
# For example ...
name = "Josuke"

print(f"{greeting}, {name}!")

# f-strings are f-ing convenient!


# let's look at defining functions in Python..

def greet(name):
    print(f"Howdy, {name}!")

# here's how we call it...

greet("partner")

# let's add a description of the function...

def greet(name):
    """
    Prints a greeting given some name.
    
    :param name: the name to be addressed in the greeting
    :type name: str
    
    """
    print(f"Howdy, {name}!")
    
# I encourage you to use docstrings!

# Python introduced support for optional type hints in v3.5.
# You can read more aobut this feature here: https://docs.python.org/3.8/library/typing.html
# let's give it a try...
def add_six(num: int) -> int:
    return num + 6

# this should print 13
print(add_six(7))

# Python also has "anonymous functions" (also known as "lambda" functions)
# take a look at the following code:

greet_alt = lambda name: print(f"Hi, {name}!")

greet_alt("Fred")

# lambda functions are often passed to other functions
# For example, they can be used to specify how a sequence should be sorted
# let's sort a list of pairs by their second element
pairs = [("bounce", 32), ("bighorn", 12), ("radical", 4), ("analysis", 7)]
# -1 is last thing in some sequence, -2 is the second to last thing in some seq, etc.
print(sorted(pairs, key=lambda pair: pair[-1]))

# we can sort it by the first element instead
# NOTE: python indexing is zero-based
print(sorted(pairs, key=lambda pair: pair[0]))

# You can learn more about other core data types and their methods here: 
# https://docs.python.org/3.8/library/stdtypes.html

# Because of its extensive standard library, Python is often described as coming with "batteries included".  
# Take a look at these "batteries": https://docs.python.org/3.8/library/

# You now know enough to complete this homework assignment (or at least where to look)
```

Output:
```
h
e
l
l
o
[2, 3, 4, 5]
[2]
hello, Josuke!
Howdy, partner!
13
Hi, Fred!
[('radical', 4), ('analysis', 7), ('bighorn', 12), ('bounce', 32)]
[('analysis', 7), ('bighorn', 12), ('bounce', 32), ('radical', 4)]
```

# Introduce yourself

Congratulations on successfully launching the notebook!  That was more than half of the challenge for this assignment. 

Please briefly introduce yourself below.  

- How would you like your instructor to refer to you?
- Where are you from?  
- What do you hope to learn in this class?  
- What are you career goals?
- Anything else you'd like to share?

<!-- grade_id: greetings, points: 0, grade: true, solution: true -->
- Wade
- Northern Virginia
- Really hoping to 
- Like to get an industry job working with all things LLMs
- Looking forward to the course!

NOTE: several of the cells in this notebook are **read only** to ensure instructions aren't unintentionally altered.  

If you can't edit the cell, it is probably intentional.

```python
from typing import *
```

# Implement `multiply(a, b)`

To test your understanding of Python basics, implement the following function which should return the product of its parameters `a` and `b`.  Add you solution after the comment **YOUR CODE HERE** and remove the `raise NotImplementedError`.  

NOTE: _Don't overthink this. It's meant to be easy!_

<!-- grade_id: multiply-problem, points: 0, grade: false, solution: true -->
```python
def multiply(a: int, b: int) -> int:
    """
    Takes two ints and returns their product
    """
    return a * b
    raise NotImplementedError()
```

<!-- grade_id: multiply-tests, points: 3, grade: true, solution: false -->
```python
assert multiply(2, 2) == 4
assert multiply(2, 0) == 0
assert multiply(-1, -1) == 1
```

## Implement `count_distinct_suffix_matches(seq, suffix)`

Complete the following function which takes a list of strings, finds the distinct elements (types), and counts how many of those types match the provided suffix.

<!-- grade_id: count-elements-with-suffix-implementation, points: 0, grade: false, solution: true -->
```python
def count_distinct_suffix_matches(seq: Iterable[Text], suffix: Text) -> int:
    """
    takes a list of strings, 
    finds the distinct elements (types), 
    and returns a count of the elements endings with the specified suffix.
    """
    unique_strings = {s for s in seq if s.endswith(suffix)}
    return len(unique_strings)
    raise NotImplementedError()
```

<!-- grade_id: count-elements-with-suffix-tests, points: 2, grade: true, solution: false -->
```python
cdsm = count_distinct_suffix_matches

assert cdsm(("father", "mother", "father", "fatherly"), suffix="er") == 2
# an empty suffix 
assert cdsm(("father", "mother", "father", "fatherly"), suffix="") == 3
assert cdsm(("father", "mother", "fatherly"), suffix="") == 3
```

# Reverse a list and remove the first element
This is exercise is just to illustrate how you can use the test cases to direct your implementation.

Implement the function `weird_request(seq)` which takes a list of strings, reverses them, drops the element at the front, and returns what remains.

For those new to Python, [this section of the official documentation may be useful](https://docs.python.org/3.7/library/stdtypes.html#common-sequence-operations)

<!-- grade_id: weird-request-problem, points: 0, grade: false, solution: true -->
```python
def weird_request(seq):
    """
    takes a list of strings, reverses them, drops the element at the front, and returns what remains.
    """
    reversed_seq = list(reversed(seq))
    reversed_seq.pop(0)

    return reversed_seq


```

<!-- grade_id: weird-request-tests, points: 2, grade: true, solution: false -->
```python
assert weird_request([1,2,3]) == [2,1]
assert weird_request([1]) == []
```